Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# Week3 Assignment - Data retrieval and dataset preprocessing.

For this assignment, you will need the iris flowers dataset available on the course website, under Week3 files. This is a dataset of 150 observations of flowers, with petal lengths and widths, and sepal lengths and widths.

The basic goal of this assignment is to preprocess the iris dataset in preparation for a machine-learning algorithm.

In the first component, you will load the dataset into a Pandas dataframe.

In the second component, you will impute missing values in one of the columns. In this case, you will assign the average value in the column to the missing data points.

In the third component, you will create two new columns that will approximate the sepal and petal sizes: One will equal the petal length multiplied by the petal width, and the other will equal the sepal length multiplied by the sepal width.

In the fourth component, you will normalize the sepal and petal sizes.

In the fifth component, you will add a column with a boolean value representing whether a flower belongs to the setosa species.

All the exercises are designed so that the solutions will need only one or a few lines of code.

Do not hesitate to contact instuctors and TA via #week3 channel on Slack if you get stuck. Join the channel first by clicking on Channels.

## Part A. Read in the iris dataset. 

In this component you will read the iris dataset into a Pandas data frame. Make sure you download the iris.csv file from the course website. Do not download or load a different instance of the iris dataset: use the one from the course website as we have modified it for this exercise, and when we test your code we will test it against our version of the dataset. Also, do not manually modify the iris dataset.

Pay attention which directory you save the file to so that you can load it by its path.
If you prefer to load the dataset by its URL, you can do that: https://biof509.github.io/spring2019/_downloads/Iris.csv


Once downloaded, use Pandas to read the file into a data frame.

Save the data frame to a variable named iris_data.

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

# Load the iris dataset into a pandas dataframe
# Make sure to save it as "iris_data"

# YOUR CODE HERE
iris_data = pd.read_csv('iris.csv')

# You can make sure you loaded it correctly by looking at its first few rows, using the .head() function
iris_data.head()
```

Output:
```
Id  SepalLengthCm  SepalWidthCm  PetalLengthCm  PetalWidthCm      Species
0   1            5.1           3.5            1.4           0.2  Iris-setosa
1   2            4.9           3.0            1.4           0.2  Iris-setosa
2   3            4.7           3.2            1.3           0.2  Iris-setosa
3   4            4.6           3.1            1.5           0.2  Iris-setosa
4   5            5.0           3.6            1.4           0.2  Iris-setosa
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert isinstance(iris_data, pd.core.frame.DataFrame)
assert iris_data.shape == (150, 6)
assert np.isclose(iris_data["SepalLengthCm"].mean(), 5.8153)
assert iris_data.isnull().sum().sum() == 7
assert iris_data.isnull().values.any()
```

## Part B. Impute missing data.

Unfortunately it appears that there are some missing data points in the PetalLengthCm column.

To resolve this, write a function that can find the mean of the PetalLengthCm column and replace the missing data points in the PetalLengthCm column with this mean.

The function takes as input a Pandas data frame and a column name. The function will return a new Pandas data frame, with the missing data points replaced with the column mean.

Run the function on the iris_data data frame and save the new data frame in the same variable (iris_data).

```python
iris_data[['PetalLengthCm']]
```

Output:
```
PetalLengthCm
0              1.4
1              1.4
2              1.3
3              1.5
4              1.4
5              1.7
6              NaN
7              1.5
8              1.4
9              1.5
10             1.5
11             1.6
12             1.4
13             1.1
14             1.2
15             1.5
16             1.3
17             1.4
18             1.7
19             1.5
20             1.7
21             1.5
22             NaN
23             1.7
24             1.9
25             1.6
26             1.6
27             1.5
28             1.4
29             1.6
..             ...
120            5.7
121            4.9
122            6.7
123            4.9
124            5.7
125            6.0
126            4.8
127            4.9
128            5.6
129            5.8
130            6.1
131            6.4
132            5.6
133            5.1
134            5.6
135            6.1
136            5.6
137            5.5
138            4.8
139            5.4
140            5.
```

```python
from sklearn.impute import SimpleImputer
```

SI = SimpleImputer()
iris_data[['PetalLengthCm']] = SI.fit_transform(iris_data[['PetalLengthCm']])

iris_data[['PetalLengthCm']]

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
def impute_with_mean(df, column_name):
    """
    Accepts a Pandas data frame and a column name as input.
    
    Returns a new Pandas data frame with missing data points in the column
    replaced with the mean value for that column.
    """
    # YOUR CODE HERE
    SI = SimpleImputer()
    df[[column_name]] = SI.fit_transform(df[[column_name]])
    return df

# Let's impute in our data frame
iris_data = impute_with_mean(iris_data, "PetalLengthCm")
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert not iris_data.isnull().values.any()
assert iris_data["PetalLengthCm"].median() == 4.25
assert iris_data.equals(impute_with_mean(iris_data, "SepalLengthCm"))
```

## Part C. Approximate full sizes.

In this component you will create two new columns that approximate the sepal and petal sizes.

The first new column will be named sepal_size and will be equal to: SepalLengthCm \* SepalWidthCm

The second new column will be named petal_size and will be equal to: PetalLengthCm \* PetalWidthCm

Add the appropriate columns below.

<!-- grade_id: cell-8d46731c6bc218b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
iris_data.insert(0, 'sepal_size', iris_data['SepalLengthCm']*iris_data['SepalWidthCm'])
iris_data.insert(0, 'petal_size', iris_data['PetalLengthCm']*iris_data['PetalWidthCm'])
```

<!-- grade_id: cell-c7edebc7b2cd1c7b, points: 1, grade: true, solution: false -->
```python
assert iris_data.shape[1] > 6
assert "sepal_size" in iris_data.columns
assert "petal_size" in iris_data.columns
assert np.isclose(np.round(iris_data["petal_size"].sum()), 868.)
assert np.isclose(np.round(iris_data["sepal_size"].sum()), 2656.)
```

## Part D. Normalize sizes.

For some machine-learning algorithms, we need to normalize data. This is generally done so that different features are on the same scales.

In our case, sepal_size and petal_size are on different scales. Normalize them by subtracting the mean from each and dividing by the standard deviation. You can do this manually if you prefer, or you can use the scale function from sklearn.preprocessing.

Note: If you do this manually, make sure to set the standard deviation to 0 degrees of freedom (ddof=0).

Since you need to perform this operation twice, you may consider writing a function to handle this, though that is not a requirement.

Save the new columns as: sepal_size_normalized, petal_size_normalized

<!-- grade_id: cell-377e7ae405383f99, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.preprocessing import StandardScaler
```

```python
scaler = StandardScaler()
norm_sepal = scaler.fit_transform(iris_data[['sepal_size']])
iris_data.insert(0,'sepal_size_normalized', norm_sepal)
norm_petal = scaler.fit_transform(iris_data[['petal_size']])
iris_data.insert(0,'petal_size_normalized', norm_petal)
```

<!-- grade_id: cell-fd44161086b39a45, points: 1, grade: true, solution: false -->
```python
assert np.isclose(iris_data["petal_size_normalized"].mean(), 0)
assert np.isclose(iris_data["petal_size_normalized"].std(ddof=0), 1)
assert np.isclose(iris_data["sepal_size_normalized"].mean(), 0)
assert np.isclose(iris_data["sepal_size_normalized"].std(ddof=0), 1)
```

## Part E. Add a boolean column.

We are specifically interested in whether a given flower is from the setosa species.

Add a column "is_setosa" to the data frame, that is true if we are dealing with a setosa and false otherwise.

<!-- grade_id: cell-75001ba281dc95cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
setosa = iris_data['Species'] == 'Iris-setosa'
iris_data.insert(0,'is_setosa',setosa)
```

<!-- grade_id: cell-b7846271cdcb45c0, points: 1, grade: true, solution: false -->
```python
assert sum(iris_data["is_setosa"]) == 50
assert iris_data.loc[iris_data["Species"] == "Iris-setosa", "is_setosa"].all()
assert iris_data["is_setosa"].dtype == "bool"
```